# Pandas

In this series of exercises you will learn to use the library `pandas`. Pandas is a very popular library for storing and manipualting data. We will be relying on the [Python Data Science Handbook](https://jakevdp.github.io/PythonDataScienceHandbook/).

Below each of the exercises is a set of assertions that test whether you gave the correct answer.

## Series 

Start by reading the introduction of `panda` objects [here](https://jakevdp.github.io/PythonDataScienceHandbook/03.01-introducing-pandas-objects.html).

### Exercise 1

Create a `Series`-object named `earnings` containing the following *figures* and using the *sources* as its index.

<!-- grade_id: cell-08d1e678831a6cb7, points: 0, grade: false, solution: true -->
```python
from pandas import Series
import pandas as pd

earnings_sources = ["sales", "ads", "subscriptions", "donations"]
earnings_figures = [39041, 8702, 13200, 292]

# your code here

display(earnings)
```

## DataFrames

### Exercise 2

Create a `DataFrame` named `skittles` with the *columns* `amount` and `rating`, using the different colors as the *index*.

|&nbsp;      | amount | rating |
|------------|--------|--------|
| **red**    | 7      | 3      |
| **green**  | 4      | 4      |
| **blue**   | 6      | 2      |
| **purple** | 5      | 4      |
| **pink**   | 6      | 3.5    |

<!-- grade_id: cell-9c6637d8f47dcc08, points: 0, grade: false, solution: true -->
```python
from pandas import DataFrame

# your code here
display(skittles)
```

<!-- grade_id: cell-1c954e8d9484f4be, points: 1, grade: true, solution: false -->
```python
assert type(skittles) is DataFrame
assert list(skittles.index) == ["red", "green", "blue", "purple", "pink"]
assert list(skittles.columns) == ["amount", "rating"]
assert skittles.loc["red", "amount"] == 7
assert skittles.loc["blue", "rating"] == 2
```

### Exercise 3

Select the column `amount` as a series and save as `skittles_amount`.

<!-- grade_id: cell-a0639cda27f32fc7, points: 0, grade: false, solution: true -->
```python
# your code here
display(skittles_amount)
```

```python
assert type(skittles_amount) = Series
```

### Exercise 4

Calculate the mean `rating` and save as `skittles_average`.

```python
# your code here
display(skittles_average)
```

<!-- grade_id: cell-4a63780e28d4318f, points: 1, grade: true, solution: false -->
```python
float(skittles_average)
assert skittles_average == 3.3
```

### Exercise 5

Add a new column to the skittles `DataFrame` called `score`. The score of a color is equal to `amount * rating`.

<!-- grade_id: cell-44c741b6c7562337, points: 0, grade: false, solution: true -->
```python
# your code here
display(skittles)
```

<!-- grade_id: cell-8d4b31a62e056a4c, points: 1, grade: true, solution: false -->
```python
assert "score" in skittles
```

## Indexing and selection

Read the [next](https://jakevdp.github.io/PythonDataScienceHandbook/03.02-data-indexing-and-selection.html) part of the reference.

### Exercise 6

For the given `DataFrame` select only columns 'a', 'c', and 'e', and rows 10, 20, 50, 60 and store the result in the variable `frame2`.

<!-- grade_id: cell-a7305488e1ae7000, points: 0, grade: false, solution: true -->
```python
import numpy as np

frame = DataFrame(np.arange(6 * 7.).reshape((6, 7)), index=[10, 20, 30, 40, 50, 60], columns=list('abcdefg'))

# your code here
display(frame2)
```

<!-- grade_id: cell-6cf72e8c3526a88a, points: 1, grade: true, solution: false -->
```python
assert type(frame2) == DataFrame
```

### Exercise 7

Replace all values in the data frame `frame` that are *divisible by 3* with the value *0*.

<!-- grade_id: cell-a0d227a359b7a2d8, points: 0, grade: false, solution: true -->
```python
frame = DataFrame(np.arange(6 * 7.).reshape((6, 7)), index=[10, 20, 30, 40, 50, 60], columns=list('abcdefg'))
# your code here
display(frame)
```

<!-- grade_id: cell-d3178809b040bf79, points: 1, grade: true, solution: false -->
```python
assert type(frame) == DataFrame
```

### Exercise 8

Select all rows from `frame` where `a > 20`, save the result in `frame3`.

<!-- grade_id: cell-65a850d571a08e13, points: 0, grade: false, solution: true -->
```python
frame = DataFrame(np.concatenate((np.arange(2 * 7.), np.arange(4 * 7., 0, -1))).reshape((6, 7)),
                  index=[10, 20, 30, 40, 50, 60], columns=list('abcdefg'))

# your code here
display(frame3)
```

## Operating on dataframes

Read about [operations in pandas](https://jakevdp.github.io/PythonDataScienceHandbook/03.03-operations-in-pandas.html).

### Exercise 9
Compute the mean for each column of `frame` below. Store the result in a Series called `mean`.

<!-- grade_id: cell-5ec92422ebfbfaff, points: 0, grade: false, solution: true -->
```python
frame = DataFrame(np.arange(6 * 7.).reshape((6, 7)), index=[10, 20, 30, 40, 50, 60], columns=list('abcdefg'))

# your code here
display(mean)
```

<!-- grade_id: cell-803ba91586590236, points: 1, grade: true, solution: false -->
```python
assert type(mean) is Series
```

### Exercise 10

Do the same as above, but compute the mean of each row instead of column.

<!-- grade_id: cell-4ac8e1f1864080a9, points: 0, grade: false, solution: true -->
```python
frame = DataFrame(np.arange(6 * 7.).reshape((6, 7)), index=[10, 20, 30, 40, 50, 60], columns=list('abcdefg'))

# your code here
display(mean)
```

<!-- grade_id: cell-74b93d2138a2653a, points: 1, grade: true, solution: false -->
```python
assert type(mean) is Series
```

This completes this `pandas` introduction.